<!DOCTYPE html>
<!--********************************************-->
<!--*       Generated from PreTeXt source      *-->
<!--*                                          *-->
<!--*         https://pretextbook.org          *-->
<!--*                                          *-->
<!--********************************************-->
<html lang="en-US">
<head xmlns:og="http://ogp.me/ns#" xmlns:book="https://ogp.me/ns/book#">
<meta http-equiv="Content-Type" content="text/html; charset=UTF-8">
<title>Chapter Assessment</title>
<meta name="Keywords" content="Authored in PreTeXt">
<meta name="viewport" content="width=device-width, initial-scale=1.0">
<meta property="og:type" content="book">
<meta property="book:title" content="Foundations of Python Programming">
<meta property="book:author" content="Matthew Hrehirchuk">
<meta property="book:author" content="Eric Chalmers">
<meta property="book:author" content="Charlotte Curtis">
<meta property="book:author" content="Patrick Perri">
<script>var runestoneMathReady = new Promise((resolve) => window.rsMathReady = resolve);
window.MathJax = {
  tex: {
    inlineMath: [['\\(','\\)']],
    tags: "none",
    tagSide: "right",
    tagIndent: ".8em",
    packages: {'[+]': ['base', 'extpfeil', 'ams', 'amscd', 'color', 'newcommand', 'knowl']}
  },
  options: {
    ignoreHtmlClass: "tex2jax_ignore|ignore-math",
    processHtmlClass: "process-math",
  },
  chtml: {
    scale: 0.98,
    mtextInheritFont: true
  },
  loader: {
    load: ['input/asciimath', '[tex]/extpfeil', '[tex]/amscd', '[tex]/color', '[tex]/newcommand', '[pretext]/mathjaxknowl3.js'],
    paths: {pretext: "https://pretextbook.org/js/lib"},
  },
  startup: {
    pageReady() {
      return MathJax.startup.defaultPageReady().then(function () {
      console.log("in ready function");
      rsMathReady();
      }
    )}
  },
};
</script><script src="https://cdn.jsdelivr.net/npm/mathjax@3/es5/tex-chtml.js"></script><link href="https://cdnjs.cloudflare.com/ajax/libs/prism/1.26.0/themes/prism.css" rel="stylesheet">
<script src="https://cdnjs.cloudflare.com/ajax/libs/prism/1.26.0/components/prism-core.min.js"></script><script src="https://cdnjs.cloudflare.com/ajax/libs/prism/1.26.0/plugins/autoloader/prism-autoloader.min.js"></script><script src="https://unpkg.com/lunr/lunr.js"></script><script src="lunr-pretext-search-index.js" async=""></script><script src="https://pretextbook.org/js/0.2/pretext_search.js"></script><link href="https://pretextbook.org/css/0.6/pretext_search.css" rel="stylesheet" type="text/css">
<script>js_version = 0.2</script><script src="https://pretextbook.org/js/lib/jquery.min.js"></script><script src="https://pretextbook.org/js/lib/jquery.sticky.js"></script><script src="https://pretextbook.org/js/lib/jquery.espy.min.js"></script><script src="https://pretextbook.org/js/0.2/pretext.js"></script><script>miniversion=0.1</script><script src="https://pretextbook.org/js/0.2/pretext_add_on.js?x=1"></script><script src="https://pretextbook.org/js/0.2/user_preferences.js"></script><script src="https://pretextbook.org/js/lib/knowl.js"></script><!--knowl.js code controls Sage Cells within knowls--><script>sagecellEvalName='Evaluate (Sage)';
</script><link rel="preconnect" href="https://fonts.googleapis.com">
<link rel="preconnect" href="https://fonts.gstatic.com" crossorigin="">
<link href="https://fonts.googleapis.com/css2?family=Inconsolata:wght@400;700&amp;family=Noto+Serif:ital,wght@0,400;0,700;1,400;1,700&amp;family=Tinos:ital,wght@0,400;0,700;1,400;1,700&amp;display=swap" rel="stylesheet">
<link href="https://fonts.cdnfonts.com/css/dejavu-serif" rel="stylesheet">
<link href="https://fonts.googleapis.com/css2?family=Roboto+Serif:opsz,wdth,wght@8..144,50..150,100..900&amp;display=swap" rel="stylesheet">
<link href="https://fonts.googleapis.com/css2?family=Open+Sans:wdth,wght@75..100,300..800&amp;display=swap" rel="stylesheet">
<link href="https://pretextbook.org/css/0.6/pretext.css" rel="stylesheet" type="text/css">
<link href="https://pretextbook.org/css/0.6/pretext_add_on.css" rel="stylesheet" type="text/css">
<link href="https://pretextbook.org/css/0.6/shell_default.css" rel="stylesheet" type="text/css">
<link href="https://pretextbook.org/css/0.6/banner_default.css" rel="stylesheet" type="text/css">
<link href="https://pretextbook.org/css/0.6/navbar_default.css" rel="stylesheet" type="text/css">
<link href="https://pretextbook.org/css/0.6/toc_default.css" rel="stylesheet" type="text/css">
<link href="https://pretextbook.org/css/0.6/knowls_default.css" rel="stylesheet" type="text/css">
<link href="https://pretextbook.org/css/0.6/style_default.css" rel="stylesheet" type="text/css">
<link href="https://pretextbook.org/css/0.6/colors_default.css" rel="stylesheet" type="text/css">
<link href="https://pretextbook.org/css/0.6/setcolors.css" rel="stylesheet" type="text/css">
<!--** eBookCongig is necessary to configure interactive       **-->
<!--** Runestone components to run locally in reader's browser **-->
<!--** No external communication:                              **-->
<!--**     log level is 0, Runestone Services are disabled     **-->
<script type="text/javascript">
eBookConfig = {};
eBookConfig.useRunestoneServices = false;
eBookConfig.host = 'http://127.0.0.1:8000';
eBookConfig.course = 'PTX_Course_Title_Here';
eBookConfig.basecourse = 'PTX_Base_Course';
eBookConfig.isLoggedIn = false;
eBookConfig.email = '';
eBookConfig.isInstructor = false;
eBookConfig.logLevel = 0;
eBookConfig.username = '';
eBookConfig.readings = null;
eBookConfig.activities = null;
eBookConfig.downloadsEnabled = false;
eBookConfig.allow_pairs = false;
eBookConfig.enableScratchAC = false;
eBookConfig.build_info = "";
eBookConfig.python3 = null;
eBookConfig.runestone_version = '7.0.0dev4';
eBookConfig.jobehost = '';
eBookConfig.proxyuri_runs = '';
eBookConfig.proxyuri_files = '';
eBookConfig.enable_chatcodes =  false;
</script>
<!--*** Runestone Services ***-->
<script type="text/javascript" src="https://runestone.academy/cdn/runestone/7.0.0dev4/prefix-runtime.3489f84c0d76065d.bundle.js"></script><script type="text/javascript" src="https://runestone.academy/cdn/runestone/7.0.0dev4/prefix-347.193e36c07abd0197.bundle.js"></script><script type="text/javascript" src="https://runestone.academy/cdn/runestone/7.0.0dev4/prefix-runestone.e543e0dd19670ae2.bundle.js"></script><link rel="stylesheet" type="text/css" href="https://runestone.academy/cdn/runestone/7.0.0dev4/prefix-347.f9add1ca35d5ad93.css">
<link rel="stylesheet" type="text/css" href="https://runestone.academy/cdn/runestone/7.0.0dev4/prefix-runestone.931faef5361cf6e9.css">
</head>
<body id="foppff" class="pretext book ignore-math">
<a class="assistive" href="#ptx-content">Skip to main content</a><header id="ptx-masthead"><div class="ptx-banner">
<a id="logo-link" href=""></a><div class="title-container">
<h1 class="heading"><a href="fopp-ff.html"><span class="title">Foundations of Python Programming:</span> <span class="subtitle">Functions First</span></a></h1>
<p class="byline">Matthew Hrehirchuk, Eric Chalmers, Charlotte Curtis, Patrick Perri</p>
</div>
<div id="searchresultsplaceholder" style="display: none">
<button id="closesearchresults" onclick="document.getElementById('searchresultsplaceholder').style.display = 'none'; return false;">x</button><h2>Search Results: <span id="searchterms"></span>
</h2>
<div id="searchempty"><span>No results.</span></div>
<ol id="searchresults"></ol>
</div>
</div></header><nav id="ptx-navbar" class="navbar"><button class="toc-toggle button" aria-label="Show or hide table of contents"><span class="icon">☰</span><span class="name">Contents</span></button><button id="user-preferences-button" class="user-preferences-button button" title="Modify user preferences"><span id="theavatarbutton" class="name">You!</span><div id="preferences_menu_holder" class="hidden"><ol id="preferences_menu" style="font-family: 'Roboto Serif', serif;">
<li data-env="avatar" tabindex="-1">Choose avatar<div class="wrap_to_submenu"><span class="to_submenu">▻</span></div>
<ol class="hidden avatar">
<li data-val="You!" tabindex="-1">
<span id="theYou!" class="avatarcheck">✔️</span>You!</li>
<li data-val="😺" tabindex="-1">
<span id="the😺" class="avatarcheck"></span>😺</li>
<li data-val="👤" tabindex="-1">
<span id="the👤" class="avatarcheck"></span>👤</li>
<li data-val="👽" tabindex="-1">
<span id="the👽" class="avatarcheck"></span>👽</li>
<li data-val="🐶" tabindex="-1">
<span id="the🐶" class="avatarcheck"></span>🐶</li>
<li data-val="🐼" tabindex="-1">
<span id="the🐼" class="avatarcheck"></span>🐼</li>
<li data-val="🌈" tabindex="-1">
<span id="the🌈" class="avatarcheck"></span>🌈</li>
</ol>
</li>
<li data-env="fontfamily" tabindex="-1">Font family<div class="wrap_to_submenu"><span class="to_submenu">▻</span></div>
<ol class="hidden fontfamily">
<li data-val="face" data-change="OS" tabindex="-1" style="font-family: 'Open Sans'">
<span id="theOS" class="ffcheck">✔️</span><span class="name">Open Sans</span><span class="sample">AaBbCc 123 PreTeXt</span>
</li>
<li data-val="face" data-change="RS" tabindex="-1" style="font-family: 'Roboto Serif'">
<span id="theRS" class="ffcheck"></span><span class="name">Roboto Serif</span><span class="sample">AaBbCc 123 PreTeXt</span>
</li>
</ol>
</li>
<li data-env="font" tabindex="-1">Adjust font<div class="wrap_to_submenu"><span class="to_submenu">▻</span></div>
<ol class="hidden fonts">
<li>Size</li>
<li><span id="thesize">12</span></li>
<li data-val="size" data-change="-1" tabindex="-1" style="font-size: 80%">Smaller</li>
<li data-val="size" data-change="1" tabindex="-1" style="font-size: 110%">Larger</li>
<li>Width</li>
<li><span id="thewdth">100</span></li>
<li data-val="wdth" data-change="-5" tabindex="-1" style="font-variation-settings: 'wdth' 60">narrower</li>
<li data-val="wdth" data-change="5" tabindex="-1" style="font-variation-settings: 'wdth' 150">wider</li>
<li>Weight</li>
<li><span id="thewght">400</span></li>
<li data-val="wght" data-change="-50" tabindex="-1" style="font-weight: 200">thinner</li>
<li data-val="wght" data-change="50" tabindex="-1" style="font-weight: 700">heavier</li>
<li>Letter spacing</li>
<li>
<span id="thelspace">0</span><span class="byunits">/200</span>
</li>
<li data-val="lspace" data-change="-1" tabindex="-1">closer</li>
<li data-val="lspace" data-change="1" tabindex="-1">f a r t h e r</li>
<li>Word spacing</li>
<li>
<span id="thewspace">0</span><span class="byunits">/50</span>
</li>
<li data-val="wspace" data-change="-1" tabindex="-1">smaller gap </li>
<li data-val="wspace" data-change="1" tabindex="-1">larger gap</li>
<li>Line Spacing</li>
<li>
<span id="theheight">135</span><span class="byunits">/100</span>
</li>
<li data-val="height" data-change="-5" tabindex="-1" style="line-height: 1">closer<br>together</li>
<li data-val="height" data-change="5" tabindex="-1" style="line-height: 1.75">further<br>apart</li>
</ol>
</li>
<li data-env="atmosphere" tabindex="-1">Light/dark mode<div class="wrap_to_submenu"><span class="to_submenu">▻</span></div>
<ol class="hidden atmosphere">
<li data-val="default" tabindex="-1">
<span id="thedefault" class="atmospherecheck">✔️</span>default</li>
<li data-val="pastel" tabindex="-1">
<span id="thepastel" class="atmospherecheck"></span>pastel</li>
<li data-val="darktwilight" tabindex="-1">
<span id="thedarktwilight" class="atmospherecheck"></span>twilight</li>
<li data-val="dark" tabindex="-1">
<span id="thedark" class="atmospherecheck"></span>dark</li>
<li data-val="darkmidnight" tabindex="-1">
<span id="thedarkmidnight" class="atmospherecheck"></span>midnight</li>
</ol>
</li>
<li data-env="ruler" tabindex="-1">Reading ruler<div class="wrap_to_submenu"><span class="to_submenu">▻</span></div>
<ol class="hidden ruler">
<li data-val="none" tabindex="-1">
<span id="thenone" class="rulercheck">✔️</span>none</li>
<li data-val="underline" tabindex="-1">
<span id="theunderline" class="rulercheck"></span>underline</li>
<li data-val="lunderline" tabindex="-1">
<span id="thelunderline" class="rulercheck"></span>L-underline</li>
<li data-val="greybar" tabindex="-1">
<span id="thegreybar" class="rulercheck"></span>grey bar</li>
<li data-val="lightbox" tabindex="-1">
<span id="thelightbox" class="rulercheck"></span>light box</li>
<li data-val="sunrise" tabindex="-1">
<span id="thesunrise" class="rulercheck"></span>sunrise</li>
<li data-val="sunriseunderline" tabindex="-1">
<span id="thesunriseunderline" class="rulercheck"></span>sunrise underline</li>
<li class="moveQ">Motion by:</li>
<li data-val="mouse" tabindex="-1">
<span id="themouse" class="motioncheck">✔️</span>follow the mouse</li>
<li data-val="arrow" tabindex="-1">
<span id="thearrow" class="motioncheck"></span>up/down arrows - not yet</li>
<li data-val="eye" tabindex="-1">
<span id="theeye" class="motioncheck"></span>eye tracking - not yet</li>
</ol>
</li>
</ol></div></button><span class="treebuttons"><a class="previous-button button" href="chapter-4.html" title="Previous"><span class="icon">&lt;</span><span class="name">Prev</span></a><a class="up-button button" href="chapter-4.html" title="Up"><span class="icon">^</span><span class="name">Up</span></a><a class="next-button button" href="chapter-5.html" title="Next"><span class="name">Next</span><span class="icon">&gt;</span></a></span><div class="searchbox"><div class="searchwidget">
<input id="ptxsearch" type="text" name="terms" placeholder="Search" onchange="doSearch()"><button id="searchbutton" type="button" onclick="doSearch()">🔍</button>
</div></div></nav><div id="latex-macros" class="hidden-content process-math" style="display:none"><span class="process-math">\(
\newcommand{\lt}{&lt;}
\newcommand{\gt}{&gt;}
\newcommand{\amp}{&amp;}
\definecolor{fillinmathshade}{gray}{0.9}
\newcommand{\fillinmath}[1]{\mathchoice{\colorbox{fillinmathshade}{$\displaystyle     \phantom{\,#1\,}$}}{\colorbox{fillinmathshade}{$\textstyle        \phantom{\,#1\,}$}}{\colorbox{fillinmathshade}{$\scriptstyle      \phantom{\,#1\,}$}}{\colorbox{fillinmathshade}{$\scriptscriptstyle\phantom{\,#1\,}$}}}
\)</span></div>
<div class="ptx-page">
<div id="ptx-sidebar"><nav id="ptx-toc" class="depth2"><ul class="structural">
<li>
<div class="toc-item"><a href="frontmatter-1.html" class="internal"><span class="title">Front Matter</span></a></div>
<ul class="structural">
<li><div class="toc-item"><a href="front-back-matter_foreword.html" class="internal"><span class="codenumber">1</span> <span class="title">Foreword to the Functions First Edition</span></a></div></li>
<li><div class="toc-item"><a href="front-back-matter_foreword_og.html" class="internal"><span class="codenumber">2</span> <span class="title">Original Foreword</span></a></div></li>
<li><div class="toc-item"><a href="front-back-matter_preface.html" class="internal"><span class="title">Preface</span></a></div></li>
<li><div class="toc-item"><a href="front-back-matter_contributor-list.html" class="internal"><span class="title">Acknowledgements</span></a></div></li>
</ul>
</li>
<li>
<div class="toc-item"><a href="chapter-1.html" class="internal"><span class="codenumber">1</span> <span class="title">General Introduction</span></a></div>
<ul class="structural">
<li>
<div class="toc-item"><a href="general-intro_introduction-the-way-of-the-program.html" class="internal"><span class="codenumber">1.1</span> <span class="title">The Way of the Program</span></a></div>
<ul class="structural">
<li><div class="toc-item"><a href="general-intro_introduction-the-way-of-the-program.html#general-intro_learning-goals" class="internal"><span class="codenumber">1.1.1</span> <span class="title">Learning Goals</span></a></div></li>
<li><div class="toc-item"><a href="general-intro_introduction-the-way-of-the-program.html#general-intro_objectives" class="internal"><span class="codenumber">1.1.2</span> <span class="title">Objectives</span></a></div></li>
</ul>
</li>
<li><div class="toc-item"><a href="general-intro_algorithms.html" class="internal"><span class="codenumber">1.2</span> <span class="title">Algorithms</span></a></div></li>
<li><div class="toc-item"><a href="general-intro_the-python-programming-language.html" class="internal"><span class="codenumber">1.3</span> <span class="title">The Python Programming Language</span></a></div></li>
<li><div class="toc-item"><a href="general-intro_special-ways-to-execute-python-in-this-book.html" class="internal"><span class="codenumber">1.4</span> <span class="title">Special Ways to Execute Python in this Book</span></a></div></li>
<li><div class="toc-item"><a href="general-intro_more-about-programs.html" class="internal"><span class="codenumber">1.5</span> <span class="title">More About Programs</span></a></div></li>
<li><div class="toc-item"><a href="general-intro_formal-and-natural-languages.html" class="internal"><span class="codenumber">1.6</span> <span class="title">Formal and Natural Languages</span></a></div></li>
<li><div class="toc-item"><a href="general-intro_a-typical-first-program.html" class="internal"><span class="codenumber">1.7</span> <span class="title">A Typical First Program</span></a></div></li>
<li><div class="toc-item"><a href="general-intro_predict-before-you-run.html" class="internal"><span class="codenumber">1.8</span> <span class="title">Predict Before You Run!</span></a></div></li>
<li><div class="toc-item"><a href="general-intro_to-understand-a-program-change-it.html" class="internal"><span class="codenumber">1.9</span> <span class="title">To Understand a Program, Change It!</span></a></div></li>
<li><div class="toc-item"><a href="general-intro_comments.html" class="internal"><span class="codenumber">1.10</span> <span class="title">Comments</span></a></div></li>
<li>
<div class="toc-item"><a href="general-intro_glossary.html" class="internal"><span class="codenumber">1.11</span> <span class="title">Glossary</span></a></div>
<ul class="structural"><li><div class="toc-item"><a href="general-intro_glossary.html#glossary-1" class="internal"><span class="codenumber">1.11</span> <span class="title">Glossary</span></a></div></li></ul>
</li>
</ul>
</li>
<li>
<div class="toc-item"><a href="chapter-2.html" class="internal"><span class="codenumber">2</span> <span class="title">Variables, Statements, and Expressions</span></a></div>
<ul class="structural">
<li>
<div class="toc-item"><a href="simple-python-data_introduction.html" class="internal"><span class="codenumber">2.1</span> <span class="title">Introduction</span></a></div>
<ul class="structural">
<li><div class="toc-item"><a href="simple-python-data_introduction.html#simple-python-data_learning-goals" class="internal"><span class="codenumber">2.1.1</span> <span class="title">Learning Goals</span></a></div></li>
<li><div class="toc-item"><a href="simple-python-data_introduction.html#simple-python-data_objectives" class="internal"><span class="codenumber">2.1.2</span> <span class="title">Objectives</span></a></div></li>
</ul>
</li>
<li><div class="toc-item"><a href="simple-python-data_values-and-data-types.html" class="internal"><span class="codenumber">2.2</span> <span class="title">Values and Data Types</span></a></div></li>
<li><div class="toc-item"><a href="simple-python-data_operators-and-operands.html" class="internal"><span class="codenumber">2.3</span> <span class="title">Operators and Operands</span></a></div></li>
<li>
<div class="toc-item"><a href="simple-python-data_function-calls.html" class="internal"><span class="codenumber">2.4</span> <span class="title">Function Calls</span></a></div>
<ul class="structural">
<li><div class="toc-item"><a href="simple-python-data_function-calls.html#simple-python-data_function-calls-as-part-of-complex-expressions" class="internal"><span class="codenumber">2.4.1</span> <span class="title">Function calls as part of complex expressions</span></a></div></li>
<li><div class="toc-item"><a href="simple-python-data_function-calls.html#simple-python-data_functions-are-objects-parentheses-invoke-functions" class="internal"><span class="codenumber">2.4.2</span> <span class="title">Functions are objects; parentheses invoke functions</span></a></div></li>
</ul>
</li>
<li><div class="toc-item"><a href="simple-python-data_data-types.html" class="internal"><span class="codenumber">2.5</span> <span class="title">Data Types</span></a></div></li>
<li><div class="toc-item"><a href="simple-python-data_type-conversion-functions.html" class="internal"><span class="codenumber">2.6</span> <span class="title">Type conversion functions</span></a></div></li>
<li><div class="toc-item"><a href="simple-python-data_variables.html" class="internal"><span class="codenumber">2.7</span> <span class="title">Variables</span></a></div></li>
<li><div class="toc-item"><a href="simple-python-data_variable-names-and-keywords.html" class="internal"><span class="codenumber">2.8</span> <span class="title">Variable Names and Keywords</span></a></div></li>
<li><div class="toc-item"><a href="simple-python-data_choosing-the-right-variable-name.html" class="internal"><span class="codenumber">2.9</span> <span class="title">Choosing the Right Variable Name</span></a></div></li>
<li><div class="toc-item"><a href="simple-python-data_statements-and-expressions.html" class="internal"><span class="codenumber">2.10</span> <span class="title">Statements and Expressions</span></a></div></li>
<li><div class="toc-item"><a href="simple-python-data_order-of-operations.html" class="internal"><span class="codenumber">2.11</span> <span class="title">Order of Operations</span></a></div></li>
<li>
<div class="toc-item"><a href="simple-python-data_reassignment.html" class="internal"><span class="codenumber">2.12</span> <span class="title">Reassignment</span></a></div>
<ul class="structural"><li><div class="toc-item"><a href="simple-python-data_reassignment.html#simple-python-data_reassignment_developing-your-mental-model-of-how-python-evaluates" class="internal"><span class="codenumber">2.12.1</span> <span class="title">Developing your mental model of How Python Evaluates</span></a></div></li></ul>
</li>
<li><div class="toc-item"><a href="simple-python-data_updating-variables.html" class="internal"><span class="codenumber">2.13</span> <span class="title">Updating Variables</span></a></div></li>
<li><div class="toc-item"><a href="simple-python-data_hard-coding.html" class="internal"><span class="codenumber">2.14</span> <span class="title">Hard-Coding</span></a></div></li>
<li><div class="toc-item"><a href="simple-python-data_input.html" class="internal"><span class="codenumber">2.15</span> <span class="title">Input</span></a></div></li>
<li>
<div class="toc-item"><a href="simple-python-data_glossary.html" class="internal"><span class="codenumber">2.16</span> <span class="title">Glossary</span></a></div>
<ul class="structural"><li><div class="toc-item"><a href="simple-python-data_glossary.html#glossary-2" class="internal"><span class="codenumber">2.16</span> <span class="title">Glossary</span></a></div></li></ul>
</li>
<li><div class="toc-item"><a href="simple-python-data_chapter-assessment.html" class="internal"><span class="codenumber">2.17</span> <span class="title">Chapter Assessment</span></a></div></li>
</ul>
</li>
<li>
<div class="toc-item"><a href="debugging_chapter.html" class="internal"><span class="codenumber">3</span> <span class="title">Debugging</span></a></div>
<ul class="structural">
<li>
<div class="toc-item"><a href="debugging_introduction-to-debugging.html" class="internal"><span class="codenumber">3.1</span> <span class="title">Introduction to Debugging</span></a></div>
<ul class="structural">
<li><div class="toc-item"><a href="debugging_introduction-to-debugging.html#debugging_learning-goals" class="internal"><span class="codenumber">3.1.1</span> <span class="title">Learning Goals</span></a></div></li>
<li><div class="toc-item"><a href="debugging_introduction-to-debugging.html#debugging_objectives" class="internal"><span class="codenumber">3.1.2</span> <span class="title">Objectives</span></a></div></li>
</ul>
</li>
<li><div class="toc-item"><a href="debugging_three-kinds-of-errors.html" class="internal"><span class="codenumber">3.2</span> <span class="title">Three kinds of Errors</span></a></div></li>
<li><div class="toc-item"><a href="debugging_syntax-errors.html" class="internal"><span class="codenumber">3.3</span> <span class="title">Syntax errors</span></a></div></li>
<li><div class="toc-item"><a href="debugging_runtime-errors.html" class="internal"><span class="codenumber">3.4</span> <span class="title">Runtime Errors</span></a></div></li>
<li>
<div class="toc-item"><a href="debugging_semantic-errors.html" class="internal"><span class="codenumber">3.5</span> <span class="title">Semantic Errors</span></a></div>
<ul class="structural"><li><div class="toc-item"><a href="debugging_semantic-errors.html#debugging_semantic-errors_test-cases" class="internal"><span class="codenumber">3.5.1</span> <span class="title">Test Cases</span></a></div></li></ul>
</li>
<li>
<div class="toc-item"><a href="debugging_know-your-error-messages.html" class="internal"><span class="codenumber">3.6</span> <span class="title">Know Your Error Messages</span></a></div>
<ul class="structural">
<li><div class="toc-item"><a href="debugging_know-your-error-messages.html#debugging_syntaxerror" class="internal"><span class="codenumber">3.6.1</span> <span class="title">SyntaxError</span></a></div></li>
<li><div class="toc-item"><a href="debugging_know-your-error-messages.html#debugging_typeerror" class="internal"><span class="codenumber">3.6.2</span> <span class="title">TypeError</span></a></div></li>
<li><div class="toc-item"><a href="debugging_know-your-error-messages.html#debugging_nameerror" class="internal"><span class="codenumber">3.6.3</span> <span class="title">NameError</span></a></div></li>
<li><div class="toc-item"><a href="debugging_know-your-error-messages.html#debugging_valueerror" class="internal"><span class="codenumber">3.6.4</span> <span class="title">ValueError</span></a></div></li>
</ul>
</li>
<li><div class="toc-item"><a href="debugging_tips-for-debugging.html" class="internal"><span class="codenumber">3.7</span> <span class="title">Tips for Debugging</span></a></div></li>
</ul>
</li>
<li>
<div class="toc-item"><a href="chapter-4.html" class="internal"><span class="codenumber">4</span> <span class="title">Functions</span></a></div>
<ul class="structural">
<li>
<div class="toc-item"><a href="functions_introduction-to-functions.html" class="internal"><span class="codenumber">4.1</span> <span class="title">Introduction to Functions</span></a></div>
<ul class="structural">
<li><div class="toc-item"><a href="functions_introduction-to-functions.html#functions_introduction-to-functions_topics" class="internal"><span class="codenumber">4.1.1</span> <span class="title">Topics</span></a></div></li>
<li><div class="toc-item"><a href="functions_introduction-to-functions.html#functions_introduction-to-functions_learning-objectives" class="internal"><span class="codenumber">4.1.2</span> <span class="title">Learning Objectives</span></a></div></li>
</ul>
</li>
<li><div class="toc-item"><a href="functions_function-definition.html" class="internal"><span class="codenumber">4.2</span> <span class="title">Function Definition</span></a></div></li>
<li><div class="toc-item"><a href="functions_function-invocation.html" class="internal"><span class="codenumber">4.3</span> <span class="title">Function Invocation</span></a></div></li>
<li><div class="toc-item"><a href="functions_function-parameters.html" class="internal"><span class="codenumber">4.4</span> <span class="title">Function Parameters</span></a></div></li>
<li><div class="toc-item"><a href="functions_returning-a-value-from-a-function.html" class="internal"><span class="codenumber">4.5</span> <span class="title">Returning a value from a function</span></a></div></li>
<li><div class="toc-item"><a href="functions_decoding-a-function.html" class="internal"><span class="codenumber">4.6</span> <span class="title">Decoding a Function</span></a></div></li>
<li><div class="toc-item"><a href="functions_type-annotations.html" class="internal"><span class="codenumber">4.7</span> <span class="title">Type Annotations</span></a></div></li>
<li><div class="toc-item"><a href="functions_variables-and-parameters-are-local.html" class="internal"><span class="codenumber">4.8</span> <span class="title">Variables and parameters are local</span></a></div></li>
<li><div class="toc-item"><a href="functions_global-variables.html" class="internal"><span class="codenumber">4.9</span> <span class="title">Global Variables</span></a></div></li>
<li><div class="toc-item"><a href="functions_functions-can-call-other-functions-composition.html" class="internal"><span class="codenumber">4.10</span> <span class="title">Functions can call other functions (Composition)</span></a></div></li>
<li><div class="toc-item"><a href="functions_flow-of-execution-summary.html" class="internal"><span class="codenumber">4.11</span> <span class="title">Flow of Execution Summary</span></a></div></li>
<li><div class="toc-item"><a href="functions_print-vs-return.html" class="internal"><span class="codenumber">4.12</span> <span class="title">Print vs. return</span></a></div></li>
<li>
<div class="toc-item"><a href="functions_glossary.html" class="internal"><span class="codenumber">4.13</span> <span class="title">Glossary</span></a></div>
<ul class="structural"><li><div class="toc-item"><a href="functions_glossary.html#glossary-3" class="internal"><span class="codenumber">4.13</span> <span class="title">Glossary</span></a></div></li></ul>
</li>
<li class="active"><div class="toc-item"><a href="functions_chapter-assessment.html" class="internal"><span class="codenumber">4.14</span> <span class="title">Chapter Assessment</span></a></div></li>
</ul>
</li>
<li>
<div class="toc-item"><a href="chapter-5.html" class="internal"><span class="codenumber">5</span> <span class="title">Conditionals</span></a></div>
<ul class="structural">
<li>
<div class="toc-item"><a href="conditionals_intro-what-we-can-do-with-turtles-and-conditionals.html" class="internal"><span class="codenumber">5.1</span> <span class="title">Intro: What we can do with Turtles and Conditionals</span></a></div>
<ul class="structural">
<li><div class="toc-item"><a href="conditionals_intro-what-we-can-do-with-turtles-and-conditionals.html#conditionals_learning-goals" class="internal"><span class="codenumber">5.1.1</span> <span class="title">Learning Goals</span></a></div></li>
<li><div class="toc-item"><a href="conditionals_intro-what-we-can-do-with-turtles-and-conditionals.html#conditionals_objectives" class="internal"><span class="codenumber">5.1.2</span> <span class="title">Objectives</span></a></div></li>
</ul>
</li>
<li><div class="toc-item"><a href="conditionals_boolean-values-and-boolean-expressions.html" class="internal"><span class="codenumber">5.2</span> <span class="title">Boolean Values and Boolean Expressions</span></a></div></li>
<li>
<div class="toc-item"><a href="conditionals_logical-operators.html" class="internal"><span class="codenumber">5.3</span> <span class="title">Logical operators</span></a></div>
<ul class="structural"><li><div class="toc-item"><a href="conditionals_logical-operators.html#conditionals_smart-evaluation" class="internal"><span class="codenumber">5.3.1</span> <span class="title">Smart Evaluation</span></a></div></li></ul>
</li>
<li><div class="toc-item"><a href="conditionals_the-in-and-not-in-operators.html" class="internal"><span class="codenumber">5.4</span> <span class="title">The <code class="code-inline tex2jax_ignore">in</code> and <code class="code-inline tex2jax_ignore">not in</code> operators</span></a></div></li>
<li><div class="toc-item"><a href="conditionals_precedence-of-operators.html" class="internal"><span class="codenumber">5.5</span> <span class="title">Precedence of Operators</span></a></div></li>
<li><div class="toc-item"><a href="conditionals_conditional-execution-binary-selection.html" class="internal"><span class="codenumber">5.6</span> <span class="title">Conditional Execution: Binary Selection</span></a></div></li>
<li><div class="toc-item"><a href="conditionals_omitting-the-else-clause-unary-selection.html" class="internal"><span class="codenumber">5.7</span> <span class="title">Omitting the <code class="code-inline tex2jax_ignore">else</code> Clause: Unary Selection</span></a></div></li>
<li><div class="toc-item"><a href="conditionals_nested-conditionals.html" class="internal"><span class="codenumber">5.8</span> <span class="title">Nested conditionals</span></a></div></li>
<li><div class="toc-item"><a href="conditionals_chained-conditionals.html" class="internal"><span class="codenumber">5.9</span> <span class="title">Chained conditionals</span></a></div></li>
<li>
<div class="toc-item"><a href="conditionals_setting-up-conditionals.html" class="internal"><span class="codenumber">5.10</span> <span class="title">Setting Up Conditionals</span></a></div>
<ul class="structural"><li><div class="toc-item"><a href="conditionals_setting-up-conditionals.html#conditionals_choosing-your-type-of-conditional" class="internal"><span class="codenumber">5.10.1</span> <span class="title">Choosing your type of Conditional</span></a></div></li></ul>
</li>
<li>
<div class="toc-item"><a href="conditionals_glossary.html" class="internal"><span class="codenumber">5.11</span> <span class="title">Glossary</span></a></div>
<ul class="structural"><li><div class="toc-item"><a href="conditionals_glossary.html#glossary-4" class="internal"><span class="codenumber">5.11</span> <span class="title">Glossary</span></a></div></li></ul>
</li>
</ul>
</li>
<li>
<div class="toc-item"><a href="chapter-6.html" class="internal"><span class="codenumber">6</span> <span class="title">Iteration</span></a></div>
<ul class="structural">
<li>
<div class="toc-item"><a href="iteration_introduction.html" class="internal"><span class="codenumber">6.1</span> <span class="title">Introduction</span></a></div>
<ul class="structural">
<li><div class="toc-item"><a href="iteration_introduction.html#iteration_learning-goals" class="internal"><span class="codenumber">6.1.1</span> <span class="title">Learning Goals</span></a></div></li>
<li><div class="toc-item"><a href="iteration_introduction.html#iteration_objectives" class="internal"><span class="codenumber">6.1.2</span> <span class="title">Objectives</span></a></div></li>
</ul>
</li>
<li><div class="toc-item"><a href="more-about-iteration_the-while-statement.html" class="internal"><span class="codenumber">6.2</span> <span class="title">The <code class="code-inline tex2jax_ignore">while</code> Statement</span></a></div></li>
<li>
<div class="toc-item"><a href="iteration_the-listener-loop.html" class="internal"><span class="codenumber">6.3</span> <span class="title">The Listener Loop</span></a></div>
<ul class="structural"><li>
<div class="toc-item"><a href="iteration_the-listener-loop.html#iteration_other-uses-of-while" class="internal"><span class="codenumber">6.3.1</span> <span class="title">Other uses of <code class="code-inline tex2jax_ignore">while</code></span></a></div>
<ul class="structural">
<li><div class="toc-item"><a href="iteration_the-listener-loop.html#iteration_sentinel-values" class="internal"><span class="codenumber">6.3.1.1</span> <span class="title">Sentinel Values</span></a></div></li>
<li><div class="toc-item"><a href="iteration_the-listener-loop.html#iteration_validating-input" class="internal"><span class="codenumber">6.3.1.2</span> <span class="title">Validating Input</span></a></div></li>
</ul>
</li></ul>
</li>
<li><div class="toc-item"><a href="iteration_break-and-continue.html" class="internal"><span class="codenumber">6.4</span> <span class="title">Break and Continue</span></a></div></li>
<li><div class="toc-item"><a href="iteration_infinite-loops.html" class="internal"><span class="codenumber">6.5</span> <span class="title">Infinite Loops</span></a></div></li>
<li><div class="toc-item"><a href="iteration_chapter-assessment.html" class="internal"><span class="codenumber">6.6</span> <span class="title">Chapter Assessment</span></a></div></li>
</ul>
</li>
<li>
<div class="toc-item"><a href="chapter-7.html" class="internal"><span class="codenumber">7</span> <span class="title">More About Iteration</span></a></div>
<ul class="structural">
<li><div class="toc-item"><a href="more-about-iteration_introduction-iteration.html" class="internal"><span class="codenumber">7.1</span> <span class="title">Introduction: Iteration</span></a></div></li>
<li><div class="toc-item"><a href="more-about-iteration_the-for-loop.html" class="internal"><span class="codenumber">7.2</span> <span class="title">The for loop</span></a></div></li>
<li><div class="toc-item"><a href="more-about-iteration_flow-of-execution-of-the-for-loop.html" class="internal"><span class="codenumber">7.3</span> <span class="title">Flow of Execution of the for Loop</span></a></div></li>
<li><div class="toc-item"><a href="more-about-iteration_strings-and-for-loops.html" class="internal"><span class="codenumber">7.4</span> <span class="title">Strings and <code class="code-inline tex2jax_ignore">for</code> loops</span></a></div></li>
<li><div class="toc-item"><a href="more-about-iteration_the-accumulator-pattern.html" class="internal"><span class="codenumber">7.5</span> <span class="title">The Accumulator Pattern</span></a></div></li>
<li><div class="toc-item"><a href="more-about-iteration_printing-intermediate-results.html" class="internal"><span class="codenumber">7.6</span> <span class="title">Printing Intermediate Results</span></a></div></li>
<li><div class="toc-item"><a href="more-about-iteration_naming-variables-in-for-loops.html" class="internal"><span class="codenumber">7.7</span> <span class="title">Naming Variables in For Loops</span></a></div></li>
<li>
<div class="toc-item"><a href="more-about-iteration_glossary.html" class="internal"><span class="codenumber">7.8</span> <span class="title">Glossary</span></a></div>
<ul class="structural"><li><div class="toc-item"><a href="more-about-iteration_glossary.html#glossary-5" class="internal"><span class="codenumber">7.8</span> <span class="title">Glossary</span></a></div></li></ul>
</li>
</ul>
</li>
<li>
<div class="toc-item"><a href="chapter-8.html" class="internal"><span class="codenumber">8</span> <span class="title">Sequences</span></a></div>
<ul class="structural">
<li>
<div class="toc-item"><a href="sequences_introduction-sequences.html" class="internal"><span class="codenumber">8.1</span> <span class="title">Introduction: Sequences</span></a></div>
<ul class="structural">
<li><div class="toc-item"><a href="sequences_introduction-sequences.html#sequences_introduction-sequences_learning-goals" class="internal"><span class="codenumber">8.1.1</span> <span class="title">Learning Goals</span></a></div></li>
<li><div class="toc-item"><a href="sequences_introduction-sequences.html#sequences_introduction-sequences_objectives" class="internal"><span class="codenumber">8.1.2</span> <span class="title">Objectives</span></a></div></li>
</ul>
</li>
<li>
<div class="toc-item"><a href="sequences_strings-and-lists.html" class="internal"><span class="codenumber">8.2</span> <span class="title">Strings and Lists</span></a></div>
<ul class="structural">
<li><div class="toc-item"><a href="sequences_strings-and-lists.html#sequences_strings-and-lists_strings" class="internal"><span class="codenumber">8.2.1</span> <span class="title">Strings</span></a></div></li>
<li><div class="toc-item"><a href="sequences_strings-and-lists.html#sequences_strings-and-lists_lists" class="internal"><span class="codenumber">8.2.2</span> <span class="title">Lists</span></a></div></li>
<li><div class="toc-item"><a href="sequences_strings-and-lists.html#sequences_strings-and-lists_tuples" class="internal"><span class="codenumber">8.2.3</span> <span class="title">Tuples</span></a></div></li>
</ul>
</li>
<li>
<div class="toc-item"><a href="sequences_index-operator-working-with-the-characters-of-a-string.html" class="internal"><span class="codenumber">8.3</span> <span class="title">Index Operator: Working with the Characters of a String</span></a></div>
<ul class="structural"><li><div class="toc-item"><a href="sequences_index-operator-working-with-the-characters-of-a-string.html#sequences_index-operator-accessing-elements-of-a-list-or-tuple" class="internal"><span class="codenumber">8.3.1</span> <span class="title">Index Operator: Accessing Elements of a List or Tuple</span></a></div></li></ul>
</li>
<li><div class="toc-item"><a href="sequences_disambiguating-creation-vs-indexing.html" class="internal"><span class="codenumber">8.4</span> <span class="title">Disambiguating []: creation vs indexing</span></a></div></li>
<li><div class="toc-item"><a href="sequences_length.html" class="internal"><span class="codenumber">8.5</span> <span class="title">Length</span></a></div></li>
<li>
<div class="toc-item"><a href="sequences_the-slice-operator.html" class="internal"><span class="codenumber">8.6</span> <span class="title">The Slice Operator</span></a></div>
<ul class="structural">
<li><div class="toc-item"><a href="sequences_the-slice-operator.html#sequences_list-slices" class="internal"><span class="codenumber">8.6.1</span> <span class="title">List Slices</span></a></div></li>
<li><div class="toc-item"><a href="sequences_the-slice-operator.html#sequences_tuple-slices" class="internal"><span class="codenumber">8.6.2</span> <span class="title">Tuple Slices</span></a></div></li>
</ul>
</li>
<li><div class="toc-item"><a href="sequences_concatenation-and-repetition.html" class="internal"><span class="codenumber">8.7</span> <span class="title">Concatenation and Repetition</span></a></div></li>
<li>
<div class="toc-item"><a href="sequences_count-and-index.html" class="internal"><span class="codenumber">8.8</span> <span class="title">Count and Index</span></a></div>
<ul class="structural">
<li><div class="toc-item"><a href="sequences_count-and-index.html#sequences_count" class="internal"><span class="codenumber">8.8.1</span> <span class="title">Count</span></a></div></li>
<li><div class="toc-item"><a href="sequences_count-and-index.html#sequences_index" class="internal"><span class="codenumber">8.8.2</span> <span class="title">Index</span></a></div></li>
</ul>
</li>
<li><div class="toc-item"><a href="sequences_splitting-and-joining-strings.html" class="internal"><span class="codenumber">8.9</span> <span class="title">Splitting and Joining Strings</span></a></div></li>
<li><div class="toc-item"><a href="sequences_chapter-assessment.html" class="internal"><span class="codenumber">8.10</span> <span class="title">Chapter Assessment</span></a></div></li>
</ul>
</li>
<li>
<div class="toc-item"><a href="chapter-9.html" class="internal"><span class="codenumber">9</span> <span class="title">Transforming Sequences</span></a></div>
<ul class="structural">
<li>
<div class="toc-item"><a href="transforming-sequences_introduction-transforming-sequences.html" class="internal"><span class="codenumber">9.1</span> <span class="title">Introduction: Transforming Sequences</span></a></div>
<ul class="structural">
<li><div class="toc-item"><a href="transforming-sequences_introduction-transforming-sequences.html#transforming-sequences_learning-goals" class="internal"><span class="codenumber">9.1.1</span> <span class="title">Learning Goals</span></a></div></li>
<li><div class="toc-item"><a href="transforming-sequences_introduction-transforming-sequences.html#transforming-sequences_objectives" class="internal"><span class="codenumber">9.1.2</span> <span class="title">Objectives</span></a></div></li>
</ul>
</li>
<li>
<div class="toc-item"><a href="transforming-sequences_mutability.html" class="internal"><span class="codenumber">9.2</span> <span class="title">Mutability</span></a></div>
<ul class="structural">
<li><div class="toc-item"><a href="transforming-sequences_mutability.html#transforming-sequences_lists-are-mutable" class="internal"><span class="codenumber">9.2.1</span> <span class="title">Lists are Mutable</span></a></div></li>
<li><div class="toc-item"><a href="transforming-sequences_mutability.html#transforming-sequences_strings-are-immutable" class="internal"><span class="codenumber">9.2.2</span> <span class="title">Strings are Immutable</span></a></div></li>
<li><div class="toc-item"><a href="transforming-sequences_mutability.html#transforming-sequences_tuples-are-immutable" class="internal"><span class="codenumber">9.2.3</span> <span class="title">Tuples are Immutable</span></a></div></li>
</ul>
</li>
<li><div class="toc-item"><a href="transforming-sequences_list-element-deletion.html" class="internal"><span class="codenumber">9.3</span> <span class="title">List Element Deletion</span></a></div></li>
<li><div class="toc-item"><a href="transforming-sequences_objects-and-references.html" class="internal"><span class="codenumber">9.4</span> <span class="title">Objects and References</span></a></div></li>
<li><div class="toc-item"><a href="transforming-sequences_aliasing.html" class="internal"><span class="codenumber">9.5</span> <span class="title">Aliasing</span></a></div></li>
<li><div class="toc-item"><a href="transforming-sequences_cloning-lists.html" class="internal"><span class="codenumber">9.6</span> <span class="title">Cloning Lists</span></a></div></li>
<li>
<div class="toc-item"><a href="transforming-sequences_mutating-methods.html" class="internal"><span class="codenumber">9.7</span> <span class="title">Mutating Methods</span></a></div>
<ul class="structural"><li><div class="toc-item"><a href="transforming-sequences_mutating-methods.html#transforming-sequences_list-methods" class="internal"><span class="codenumber">9.7.1</span> <span class="title">List Methods</span></a></div></li></ul>
</li>
<li><div class="toc-item"><a href="transforming-sequences_append-versus-concatenate.html" class="internal"><span class="codenumber">9.8</span> <span class="title">Append versus Concatenate</span></a></div></li>
<li>
<div class="toc-item"><a href="transforming-sequences_non-mutating-methods-on-strings.html" class="internal"><span class="codenumber">9.9</span> <span class="title">Non-mutating Methods on Strings</span></a></div>
<ul class="structural"><li><div class="toc-item"><a href="transforming-sequences_non-mutating-methods-on-strings.html#transforming-sequences_string-format-method" class="internal"><span class="codenumber">9.9.1</span> <span class="title">FStrings</span></a></div></li></ul>
</li>
<li><div class="toc-item"><a href="transforming-sequences_the-accumulator-pattern-with-lists.html" class="internal"><span class="codenumber">9.10</span> <span class="title">The Accumulator Pattern with Lists</span></a></div></li>
<li><div class="toc-item"><a href="transforming-sequences_the-accumulator-pattern-with-strings.html" class="internal"><span class="codenumber">9.11</span> <span class="title">The Accumulator Pattern with Strings</span></a></div></li>
<li>
<div class="toc-item"><a href="transforming-sequences_accumulator-pattern-strategies.html" class="internal"><span class="codenumber">9.12</span> <span class="title">Accumulator Pattern Strategies</span></a></div>
<ul class="structural">
<li><div class="toc-item"><a href="transforming-sequences_accumulator-pattern-strategies.html#transforming-sequences_when-to-use-it" class="internal"><span class="codenumber">9.12.1</span> <span class="title">When to Use it</span></a></div></li>
<li><div class="toc-item"><a href="transforming-sequences_accumulator-pattern-strategies.html#transforming-sequences_before-writing-it" class="internal"><span class="codenumber">9.12.2</span> <span class="title">Before Writing it</span></a></div></li>
<li><div class="toc-item"><a href="transforming-sequences_accumulator-pattern-strategies.html#transforming-sequences_choosing-good-accumulator-and-iterator-variable-names" class="internal"><span class="codenumber">9.12.3</span> <span class="title">Choosing Good Accumulator and Iterator Variable Names</span></a></div></li>
</ul>
</li>
<li><div class="toc-item"><a href="transforming-sequences_dont-mutate-a-list-that-you-are-iterating-through.html" class="internal"><span class="codenumber">9.13</span> <span class="title">Don't Mutate A List That You Are Iterating Through</span></a></div></li>
<li>
<div class="toc-item"><a href="transforming-sequences_summary.html" class="internal"><span class="codenumber">9.14</span> <span class="title">Summary</span></a></div>
<ul class="structural"><li><div class="toc-item"><a href="transforming-sequences_summary.html#glossary-6" class="internal"><span class="codenumber">9.14</span> <span class="title">Glossary</span></a></div></li></ul>
</li>
</ul>
</li>
<li>
<div class="toc-item"><a href="chapter-10.html" class="internal"><span class="codenumber">10</span> <span class="title">Nested Data and Nested Iteration</span></a></div>
<ul class="structural">
<li>
<div class="toc-item"><a href="nested-data_introduction-nested-data-and-nested-iteration.html" class="internal"><span class="codenumber">10.1</span> <span class="title">Introduction: Nested Data and Nested Iteration</span></a></div>
<ul class="structural"><li><div class="toc-item"><a href="nested-data_introduction-nested-data-and-nested-iteration.html#nested-data_lists-with-complex-items" class="internal"><span class="codenumber">10.1.1</span> <span class="title">Lists with Complex Items</span></a></div></li></ul>
</li>
<li><div class="toc-item"><a href="nested-data_nested-dictionaries.html" class="internal"><span class="codenumber">10.2</span> <span class="title">Nested Dictionaries</span></a></div></li>
<li><div class="toc-item"><a href="nested-data_processing-json-results.html" class="internal"><span class="codenumber">10.3</span> <span class="title">Processing JSON results</span></a></div></li>
<li><div class="toc-item"><a href="nested-data_nested-iteration.html" class="internal"><span class="codenumber">10.4</span> <span class="title">Nested Iteration</span></a></div></li>
<li><div class="toc-item"><a href="nested-data_structuring-nested-data.html" class="internal"><span class="codenumber">10.5</span> <span class="title">Structuring Nested Data</span></a></div></li>
<li><div class="toc-item"><a href="nested-data_deep-and-shallow-copies.html" class="internal"><span class="codenumber">10.6</span> <span class="title">Deep and Shallow Copies</span></a></div></li>
<li><div class="toc-item"><a href="nested-data_chapter-assessment.html" class="internal"><span class="codenumber">10.7</span> <span class="title">Chapter Assessment</span></a></div></li>
</ul>
</li>
<li>
<div class="toc-item"><a href="chapter-11.html" class="internal"><span class="codenumber">11</span> <span class="title">Files</span></a></div>
<ul class="structural">
<li>
<div class="toc-item"><a href="files_introduction-working-with-data-files.html" class="internal"><span class="codenumber">11.1</span> <span class="title">Introduction: Working with Data Files</span></a></div>
<ul class="structural">
<li><div class="toc-item"><a href="files_introduction-working-with-data-files.html#files_learning-goals" class="internal"><span class="codenumber">11.1.1</span> <span class="title">Learning Goals</span></a></div></li>
<li><div class="toc-item"><a href="files_introduction-working-with-data-files.html#files_objectives" class="internal"><span class="codenumber">11.1.2</span> <span class="title">Objectives</span></a></div></li>
</ul>
</li>
<li><div class="toc-item"><a href="files_reading-a-file.html" class="internal"><span class="codenumber">11.2</span> <span class="title">Reading a File</span></a></div></li>
<li><div class="toc-item"><a href="files_alternative-file-reading-methods.html" class="internal"><span class="codenumber">11.3</span> <span class="title">Alternative File Reading Methods</span></a></div></li>
<li><div class="toc-item"><a href="files_iterating-over-lines-in-a-file.html" class="internal"><span class="codenumber">11.4</span> <span class="title">Iterating over lines in a file</span></a></div></li>
<li><div class="toc-item"><a href="files_finding-a-file-in-your-filesystem.html" class="internal"><span class="codenumber">11.5</span> <span class="title">Finding a File in your Filesystem</span></a></div></li>
<li><div class="toc-item"><a href="files_using-with-for-files.html" class="internal"><span class="codenumber">11.6</span> <span class="title">Using <code class="code-inline tex2jax_ignore">with</code> for Files</span></a></div></li>
<li><div class="toc-item"><a href="files_recipe-for-reading-and-processing-a-file.html" class="internal"><span class="codenumber">11.7</span> <span class="title">Recipe for Reading and Processing a File</span></a></div></li>
<li><div class="toc-item"><a href="files_writing-text-files.html" class="internal"><span class="codenumber">11.8</span> <span class="title">Writing Text Files</span></a></div></li>
<li><div class="toc-item"><a href="files_csv-format.html" class="internal"><span class="codenumber">11.9</span> <span class="title">CSV Format</span></a></div></li>
<li><div class="toc-item"><a href="files_reading-in-data-from-a-csv-file.html" class="internal"><span class="codenumber">11.10</span> <span class="title">Reading in data from a CSV File</span></a></div></li>
<li><div class="toc-item"><a href="files_writing-data-to-a-csv-file.html" class="internal"><span class="codenumber">11.11</span> <span class="title">Writing data to a CSV File</span></a></div></li>
<li><div class="toc-item"><a href="files_tips-on-handling-files.html" class="internal"><span class="codenumber">11.12</span> <span class="title">Tips on Handling Files</span></a></div></li>
<li>
<div class="toc-item"><a href="files_glossary.html" class="internal"><span class="codenumber">11.13</span> <span class="title">Glossary</span></a></div>
<ul class="structural"><li><div class="toc-item"><a href="files_glossary.html#glossary-7" class="internal"><span class="codenumber">11.13</span> <span class="title">Glossary</span></a></div></li></ul>
</li>
<li><div class="toc-item"><a href="files_chapter-assessment.html" class="internal"><span class="codenumber">11.14</span> <span class="title">Chapter Assessment</span></a></div></li>
</ul>
</li>
<li>
<div class="toc-item"><a href="chapter-12.html" class="internal"><span class="codenumber">12</span> <span class="title">Exceptions</span></a></div>
<ul class="structural">
<li>
<div class="toc-item"><a href="exceptions_when-to-use-try-except.html" class="internal"><span class="codenumber">12.1</span> <span class="title">Handling Exceptions</span></a></div>
<ul class="structural"><li><div class="toc-item"><a href="exceptions_when-to-use-try-except.html#subsection-47" class="internal"><span class="codenumber">12.1.1</span> <span class="title">When to use try/except</span></a></div></li></ul>
</li>
<li><div class="toc-item"><a href="exceptions_standard-exceptions79_index-0.html" class="internal"><span class="codenumber">12.2</span> <span class="title">Standard Exceptions</span></a></div></li>
<li><div class="toc-item"><a href="exceptions_chapter-assessment.html" class="internal"><span class="codenumber">12.3</span> <span class="title">Chapter Assessment</span></a></div></li>
</ul>
</li>
<li>
<div class="toc-item"><a href="chapter-13.html" class="internal"><span class="codenumber">13</span> <span class="title">Defining your own Classes</span></a></div>
<ul class="structural">
<li>
<div class="toc-item"><a href="classes_introduction-classes-and-objects-the-basics.html" class="internal"><span class="codenumber">13.1</span> <span class="title">Introduction: Classes and Objects - the Basics</span></a></div>
<ul class="structural"><li><div class="toc-item"><a href="classes_introduction-classes-and-objects-the-basics.html#classes_object-oriented-programming" class="internal"><span class="codenumber">13.1.1</span> <span class="title">Object-oriented programming</span></a></div></li></ul>
</li>
<li><div class="toc-item"><a href="classes_objects-revisited.html" class="internal"><span class="codenumber">13.2</span> <span class="title">Objects Revisited</span></a></div></li>
<li><div class="toc-item"><a href="classes_user-defined-classes.html" class="internal"><span class="codenumber">13.3</span> <span class="title">User Defined Classes</span></a></div></li>
<li><div class="toc-item"><a href="classes_adding-parameters-to-the-constructor.html" class="internal"><span class="codenumber">13.4</span> <span class="title">Adding Parameters to the Constructor</span></a></div></li>
<li><div class="toc-item"><a href="classes_adding-other-methods-to-a-class.html" class="internal"><span class="codenumber">13.5</span> <span class="title">Adding Other Methods to a Class</span></a></div></li>
<li><div class="toc-item"><a href="classes_objects-as-arguments-and-parameters.html" class="internal"><span class="codenumber">13.6</span> <span class="title">Objects as Arguments and Parameters</span></a></div></li>
<li><div class="toc-item"><a href="classes_converting-an-object-to-a-string.html" class="internal"><span class="codenumber">13.7</span> <span class="title">Converting an Object to a String</span></a></div></li>
<li><div class="toc-item"><a href="classes_instances-as-return-values.html" class="internal"><span class="codenumber">13.8</span> <span class="title">Instances as Return Values</span></a></div></li>
<li><div class="toc-item"><a href="classes_class-variables-and-instance-variables.html" class="internal"><span class="codenumber">13.9</span> <span class="title">Class Variables and Instance Variables</span></a></div></li>
<li><div class="toc-item"><a href="classes_thinking-about-classes-and-instances.html" class="internal"><span class="codenumber">13.10</span> <span class="title">Thinking About Classes and Instances</span></a></div></li>
<li><div class="toc-item"><a href="classes_a-tamagotchi-game.html" class="internal"><span class="codenumber">13.11</span> <span class="title">A Tamagotchi Game</span></a></div></li>
<li>
<div class="toc-item"><a href="classes_glossary.html" class="internal"><span class="codenumber">13.12</span> <span class="title">Glossary</span></a></div>
<ul class="structural"><li><div class="toc-item"><a href="classes_glossary.html#glossary-8" class="internal"><span class="codenumber">13.12</span> <span class="title">Glossary</span></a></div></li></ul>
</li>
<li><div class="toc-item"><a href="classes_chapter-assessment.html" class="internal"><span class="codenumber">13.13</span> <span class="title">Chapter Assessment</span></a></div></li>
</ul>
</li>
<li>
<div class="toc-item"><a href="chapter-14.html" class="internal"><span class="codenumber">14</span> <span class="title">Additional Resources</span></a></div>
<ul class="structural">
<li>
<div class="toc-item"><a href="additional_resources-command_line.html" class="internal"><span class="codenumber">14.1</span> <span class="title">The Command Line</span></a></div>
<ul class="structural">
<li><div class="toc-item"><a href="additional_resources-command_line.html#command_line-git-bash" class="internal"><span class="codenumber">14.1.1</span> <span class="title">GitBash</span></a></div></li>
<li>
<div class="toc-item"><a href="additional_resources-command_line.html#command_line-commands" class="internal"><span class="codenumber">14.1.2</span> <span class="title">Commands</span></a></div>
<ul class="structural">
<li><div class="toc-item"><a href="additional_resources-command_line.html#command_line-commands-navigation" class="internal"><span class="codenumber">14.1.2.1</span> <span class="title">Navigation Commands</span></a></div></li>
<li><div class="toc-item"><a href="additional_resources-command_line.html#command_line-commands-files_directories" class="internal"><span class="codenumber">14.1.2.2</span> <span class="title">File and Directory Commands</span></a></div></li>
</ul>
</li>
</ul>
</li>
<li>
<div class="toc-item"><a href="additional_resources-intro-to-git.html" class="internal"><span class="codenumber">14.2</span> <span class="title">Intro to Git</span></a></div>
<ul class="structural">
<li><div class="toc-item"><a href="additional_resources-intro-to-git.html#intro-to-git_the-history-of-git" class="internal"><span class="codenumber">14.2.1</span> <span class="title">The History of Git</span></a></div></li>
<li><div class="toc-item"><a href="additional_resources-intro-to-git.html#intro-to-git_repositories" class="internal"><span class="codenumber">14.2.2</span> <span class="title">Repositories</span></a></div></li>
<li><div class="toc-item"><a href="additional_resources-intro-to-git.html#intro-to-git_cloning-a-repository" class="internal"><span class="codenumber">14.2.3</span> <span class="title">Cloning a Repository</span></a></div></li>
<li><div class="toc-item"><a href="additional_resources-intro-to-git.html#intro-to-git_staging-and-committing" class="internal"><span class="codenumber">14.2.4</span> <span class="title">Staging and Committing</span></a></div></li>
<li><div class="toc-item"><a href="additional_resources-intro-to-git.html#intro-to-git_pushing-and-pulling" class="internal"><span class="codenumber">14.2.5</span> <span class="title">Pushing and Pulling</span></a></div></li>
</ul>
</li>
<li>
<div class="toc-item"><a href="additional_resources-advanced-git.html" class="internal"><span class="codenumber">14.3</span> <span class="title">Advanced Git</span></a></div>
<ul class="structural">
<li><div class="toc-item"><a href="additional_resources-advanced-git.html#advanced-git_conflicts" class="internal"><span class="codenumber">14.3.1</span> <span class="title">Conflicts</span></a></div></li>
<li><div class="toc-item"><a href="additional_resources-advanced-git.html#advanced-git_branching" class="internal"><span class="codenumber">14.3.2</span> <span class="title">Branching</span></a></div></li>
<li><div class="toc-item"><a href="additional_resources-advanced-git.html#advanced-git_merging" class="internal"><span class="codenumber">14.3.3</span> <span class="title">Merging Branches</span></a></div></li>
<li><div class="toc-item"><a href="additional_resources-advanced-git.html#advanced-git_rebasing" class="internal"><span class="codenumber">14.3.4</span> <span class="title">Rebasing</span></a></div></li>
<li><div class="toc-item"><a href="additional_resources-advanced-git.html#advanced-git_stashing" class="internal"><span class="codenumber">14.3.5</span> <span class="title">Stashing</span></a></div></li>
</ul>
</li>
<li>
<div class="toc-item"><a href="additional_resources-books-bookshelves.html" class="internal"><span class="codenumber">14.4</span> <span class="title">Comprehensive Practice Program: Books and Bookshelves</span></a></div>
<ul class="structural">
<li><div class="toc-item"><a href="additional_resources-books-bookshelves.html#subsec-additonal_resources-books-bookshelves-adding-books" class="internal"><span class="codenumber">14.4.1</span> <span class="title">Adding Some Books</span></a></div></li>
<li><div class="toc-item"><a href="additional_resources-books-bookshelves.html#additonal_resources-books-bookshelves-confirming-books" class="internal"><span class="codenumber">14.4.2</span> <span class="title">Confirming Book Data</span></a></div></li>
<li><div class="toc-item"><a href="additional_resources-books-bookshelves.html#additonal_resources-books-bookshelves-searching-books" class="internal"><span class="codenumber">14.4.3</span> <span class="title">Searching for Books</span></a></div></li>
<li><div class="toc-item"><a href="additional_resources-books-bookshelves.html#additonal_resources-books-bookshelves-adding-more-books" class="internal"><span class="codenumber">14.4.4</span> <span class="title">Filling Up Bookshelves</span></a></div></li>
<li><div class="toc-item"><a href="additional_resources-books-bookshelves.html#additonal_resources-books-bookshelves-final-touches" class="internal"><span class="codenumber">14.4.5</span> <span class="title">Final Touches</span></a></div></li>
</ul>
</li>
</ul>
</li>
<li>
<div class="toc-item"><a href="backmatter-1.html" class="internal"><span class="title">Back Matter</span></a></div>
<ul class="structural">
<li>
<div class="toc-item"><a href="front-back-matter_gnu-free-documentation-license.html" class="internal"><span class="codenumber">A</span> <span class="title">GNU Free Documentation License</span></a></div>
<ul class="structural">
<li><div class="toc-item"><a href="front-back-matter_preamble.html" class="internal"><span class="codenumber">A.1</span> <span class="title">0. PREAMBLE</span></a></div></li>
<li><div class="toc-item"><a href="front-back-matter_applicability-and-definitions.html" class="internal"><span class="codenumber">A.2</span> <span class="title">1. APPLICABILITY AND DEFINITIONS</span></a></div></li>
<li><div class="toc-item"><a href="front-back-matter_verbatim-copying.html" class="internal"><span class="codenumber">A.3</span> <span class="title">2. VERBATIM COPYING</span></a></div></li>
<li><div class="toc-item"><a href="front-back-matter_copying-in-quantity.html" class="internal"><span class="codenumber">A.4</span> <span class="title">3. COPYING IN QUANTITY</span></a></div></li>
<li><div class="toc-item"><a href="front-back-matter_modifications.html" class="internal"><span class="codenumber">A.5</span> <span class="title">4. MODIFICATIONS</span></a></div></li>
<li><div class="toc-item"><a href="front-back-matter_combining-documents.html" class="internal"><span class="codenumber">A.6</span> <span class="title">5. COMBINING DOCUMENTS</span></a></div></li>
<li><div class="toc-item"><a href="front-back-matter_collections-of-documents.html" class="internal"><span class="codenumber">A.7</span> <span class="title">6. COLLECTIONS OF DOCUMENTS</span></a></div></li>
<li><div class="toc-item"><a href="front-back-matter_aggregation-with-independent-works.html" class="internal"><span class="codenumber">A.8</span> <span class="title">7. AGGREGATION WITH INDEPENDENT WORKS</span></a></div></li>
<li><div class="toc-item"><a href="front-back-matter_translation.html" class="internal"><span class="codenumber">A.9</span> <span class="title">8. TRANSLATION</span></a></div></li>
<li><div class="toc-item"><a href="front-back-matter_termination.html" class="internal"><span class="codenumber">A.10</span> <span class="title">9. TERMINATION</span></a></div></li>
<li><div class="toc-item"><a href="front-back-matter_future-revisions-of-this-license.html" class="internal"><span class="codenumber">A.11</span> <span class="title">10. FUTURE REVISIONS OF THIS LICENSE</span></a></div></li>
<li><div class="toc-item"><a href="front-back-matter_relicensing.html" class="internal"><span class="codenumber">A.12</span> <span class="title">11. RELICENSING</span></a></div></li>
</ul>
</li>
<li><div class="toc-item"><a href="front-back-matter_copyright-notice.html" class="internal"><span class="codenumber">B</span> <span class="title">Copyright Notice</span></a></div></li>
</ul>
</li>
</ul></nav></div>
<main class="ptx-main"><div id="ptx-content" class="ptx-content"><section class="section" id="functions_chapter-assessment"><h2 class="heading hide-type">
<span class="type">Section</span> <span class="codenumber">4.14</span> <span class="title">Chapter Assessment</span>
</h2>
<article class="exercise exercise-like" id="ac4_15_1"><h3 class="heading">
<span class="type">Checkpoint</span><span class="space"> </span><span class="codenumber">4.14.1</span><span class="period">.</span>
</h3>
<div class="ptx-runestone-container"><div class="runestone explainer ac_section "><div data-component="activecode" id="ac4_15_1">
<div class="ac_question col-md-12" id="ac4_15_1_question"><div class="para" id="p-1359">Write a function called <code class="code-inline tex2jax_ignore">int_return</code> that takes an integer as input and returns the same integer.</div></div>
<textarea data-lang="python" data-timelimit="25000" data-audio="" data-coach="true" style="visibility: hidden;" id="ac4_15_1_editor" data-question_label="" data-codelens="true">=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

  def testOne(self):
    self.assertEqual(int_return(10), 10, "Testing that function int_return(10) returns 10")
myTests().main()
</textarea>
</div></div></div></article><article class="exercise exercise-like" id="ac4_15_2"><h3 class="heading">
<span class="type">Checkpoint</span><span class="space"> </span><span class="codenumber">4.14.2</span><span class="period">.</span>
</h3>
<div class="ptx-runestone-container"><div class="runestone explainer ac_section "><div data-component="activecode" id="ac4_15_2">
<div class="ac_question col-md-12" id="ac4_15_2_question"><div class="para" id="p-1360">Write a function called <code class="code-inline tex2jax_ignore">add</code> that takes any number as its input and returns that sum with 2 added.</div></div>
<textarea data-lang="python" data-timelimit="25000" data-audio="" data-coach="true" style="visibility: hidden;" id="ac4_15_2_editor" data-question_label="" data-codelens="true">=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

  def testTwo(self):
    self.assertEqual(add(-2), 0, "Testing that add(-2) returns 0")
    self.assertEqual(add(6), 8, "Testing that add(6) returns 8")
    self.assertEqual(add(4), 6, "Testing that add(4) returns 6")

myTests().main()
</textarea>
</div></div></div></article><article class="exercise exercise-like" id="ac4_15_3"><h3 class="heading">
<span class="type">Checkpoint</span><span class="space"> </span><span class="codenumber">4.14.3</span><span class="period">.</span>
</h3>
<div class="ptx-runestone-container"><div class="runestone explainer ac_section "><div data-component="activecode" id="ac4_15_3">
<div class="ac_question col-md-12" id="ac4_15_3_question"><div class="para" id="p-1361">Write a function called <code class="code-inline tex2jax_ignore">change</code> that takes any string, adds “Nice to meet you!” to the end of the argument given, and returns that new string.</div></div>
<textarea data-lang="python" data-timelimit="25000" data-audio="" data-coach="true" style="visibility: hidden;" id="ac4_15_3_editor" data-question_label="" data-codelens="true">=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

  def testThree(self):
    self.assertEqual(change("I'm Bob. "), "I'm Bob. Nice to meet you!", "Tests that change('I'm Bob. '') returns 'I'm Bob. Nice to meet you!'")
    self.assertEqual(change(""), "Nice to meet you!", "Tests that change() returns 'Nice to meet you!'")

myTests().main()
</textarea>
</div></div></div></article><article class="exercise exercise-like" id="ac4_15_4"><h3 class="heading">
<span class="type">Checkpoint</span><span class="space"> </span><span class="codenumber">4.14.4</span><span class="period">.</span>
</h3>
<div class="ptx-runestone-container"><div class="runestone explainer ac_section "><div data-component="activecode" id="ac4_15_4">
<div class="ac_question col-md-12" id="ac4_15_4_question"><div class="para" id="p-1362">Write a function, <code class="code-inline tex2jax_ignore">addup</code>, that takes five integers as inputs and returns the sum of those integers.</div></div>
<textarea data-lang="python" data-timelimit="25000" data-audio="" data-coach="true" style="visibility: hidden;" id="ac4_15_4_editor" data-question_label="" data-codelens="true">=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

  def testFourA(self):
    self.assertEqual(addup(1, 2, 3, 4, 5), 15, "Tests that addup(1, 2, 3, 4, 5) returns 15")
    self.assertEqual(addup(0, 0, 0, 0, 1), 1, "Tests that addup(0, 0, 0, 0, 1) returns 1")
    self.assertEqual(addup(-1, 1, -2, 2, -3), -3, "Tests that addup(-1, 1, -2, 2, -3) returns -3")
myTests().main()
</textarea>
</div></div></div></article><article class="exercise exercise-like" id="ac4_15_5"><h3 class="heading">
<span class="type">Checkpoint</span><span class="space"> </span><span class="codenumber">4.14.5</span><span class="period">.</span>
</h3>
<div class="ptx-runestone-container"><div class="runestone explainer ac_section "><div data-component="activecode" id="ac4_15_5">
<div class="ac_question col-md-12" id="ac4_15_5_question"><div class="para" id="p-1363">You will need to write two functions for this problem. The first function, <code class="code-inline tex2jax_ignore">divide</code> that takes in any number and returns that same number divided by 2. The second function called <code class="code-inline tex2jax_ignore">sum</code> should take any number, divide it by 2, and add 6. It should return this new number. You should call the <code class="code-inline tex2jax_ignore">divide</code> function within the <code class="code-inline tex2jax_ignore">sum</code> function. Do not worry about decimals.</div></div>
<textarea data-lang="python" data-timelimit="25000" data-audio="" data-coach="true" style="visibility: hidden;" id="ac4_15_5_editor" data-question_label="" data-codelens="true">=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

  def testSixA(self):
    self.assertEqual(divide(4), 2, "Tests that divide(4) returns 2")
  def testSixB(self):
    self.assertEqual(sum(4), 8, "Tests that sum(4) returns 8")
    self.assertEqual(sum(2), 7, "Tests that sum(2) returns 7")
    self.assertEqual(sum(-6), 3, "Tests that sum(-6) returns 3")
    self.assertEqual(sum(0), 6, "Tests that sum(0) returns 6")

myTests().main()
</textarea>
</div></div></div></article><article class="exercise exercise-like" id="ac4_15_6"><h3 class="heading">
<span class="type">Checkpoint</span><span class="space"> </span><span class="codenumber">4.14.6</span><span class="period">.</span>
</h3>
<div class="ptx-runestone-container"><div class="runestone explainer ac_section "><div data-component="activecode" id="ac4_15_6">
<div class="ac_question col-md-12" id="ac4_15_6_question">
<div class="para" id="p-1364">You will need to write three functions for this problem.</div> <ul class="disc">
<li id="li-207"><div class="para" id="p-1365">The first function, <code class="code-inline tex2jax_ignore">average</code> that takes in any three numbers and returns the average.</div></li>
<li id="li-208"><div class="para" id="p-1366">The second function called <code class="code-inline tex2jax_ignore">diffSquare</code> that takes one number and the average of the three numbers, subtracts those numbers, and then returns the square of this difference.</div></li>
<li id="li-209">
<div class="para" id="p-1367">A third function called <code class="code-inline tex2jax_ignore">stdev</code> which is passed the three numbers and then:</div>
<ol class="decimal">
<li id="li-210"><div class="para" id="p-1368">Calls on <code class="code-inline tex2jax_ignore">average</code> to find the averaage of those three values, and assign the resulting average to a local variable.</div></li>
<li id="li-211"><div class="para" id="p-1369">Calls on <code class="code-inline tex2jax_ignore">diffSquare</code> three times, passing each of the three numbers in turn and the local variable average, assigning the resulting squared differences to three local variables.</div></li>
<li id="li-212"><div class="para" id="p-1370">Calls on <code class="code-inline tex2jax_ignore">average</code> again, this time passing in the three squared differences, and assigning this new average to a variable called <code class="code-inline tex2jax_ignore">variance</code>.</div></li>
</ol>
</li>
<li id="li-213"><div class="para" id="p-1371">Finally, <code class="code-inline tex2jax_ignore">stdev</code> finds the square root of <code class="code-inline tex2jax_ignore">variance</code> as the standard deviation</div></li>
</ul> <div class="para" id="p-1372">Do not worry about rounding the decimals.</div>
</div>
<textarea data-lang="python" data-timelimit="25000" data-audio="" data-coach="true" style="visibility: hidden;" id="ac4_15_6_editor" data-question_label="" data-codelens="true">"""
Using these values: 10, 13, 15,
you should get:
- an average of 12.66666667.
- these squared differences are: 7.1111111 0.11111111, 5.4444444
- which sum to 12.66666667 and average to be 4.2222222
- the square root of which is: 2.054804667
""""

def average( ,  ,  ): 




    return mean


def diffSquare(   ,   ): 





    return squaredDifference


def stdev(   ,   ,   ):


    variance = average (  ,   ,   )
    return variance**(1/2)

print('The standard deviation of the values 10, 13 and 15 is', stdev(10,13,15))
</textarea>
</div></div></div></article></section></div>
<div class="ptx-content-footer">
<a class="previous-button button" href="chapter-4.html" title="Previous"><span class="icon">&lt;</span><span class="name">Prev</span></a><a class="top-button button" href="#" title="Top"><span class="icon">^</span><span class="name">Top</span></a><a class="next-button button" href="chapter-5.html" title="Next"><span class="name">Next</span><span class="icon">&gt;</span></a>
</div></main>
</div>
<div id="ptx-page-footer">
<a class="pretext-link" href="https://pretextbook.org" title="PreTeXt"><div class="logo"><svg xmlns="http://www.w3.org/2000/svg" width="100%" height="100%" viewBox="338 3000 8772 6866"><g style="stroke-width:.025in; stroke:black; fill:none"><polyline points="472,3590 472,9732 " style="stroke:#000000;stroke-width:174; stroke-linejoin:miter; stroke-linecap:round; "></polyline><path style="stroke:#000000;stroke-width:126;stroke-linecap:butt;" d="M 4724,9448 A 4660 4660  0  0  1  8598  9259"></path><path style="stroke:#000000;stroke-width:174;stroke-linecap:butt;" d="M 4488,9685 A 4228 4228  0  0  0   472  9732"></path><path style="stroke:#000000;stroke-width:126;stroke-linecap:butt;" d="M 4724,3590 A 4241 4241  0  0  1  8598  3496"></path><path style="stroke:#000000;stroke-width:126;stroke-linecap:round;" d="M 850,3496  A 4241 4241  0  0  1  4724  3590"></path><path style="stroke:#000000;stroke-width:126;stroke-linecap:round;" d="M 850,9259  A 4507 4507  0  0  1  4724  9448"></path><polyline points="5385,4299 4062,8125" style="stroke:#000000;stroke-width:300; stroke-linejoin:miter; stroke-linecap:round;"></polyline><polyline points="8598,3496 8598,9259" style="stroke:#000000;stroke-width:126; stroke-linejoin:miter; stroke-linecap:round;"></polyline><polyline points="850,3496 850,9259" style="stroke:#000000;stroke-width:126; stroke-linejoin:miter; stroke-linecap:round;"></polyline><polyline points="4960,9685 4488,9685" style="stroke:#000000;stroke-width:174; stroke-linejoin:miter; stroke-linecap:round;"></polyline><polyline points="3070,4582 1889,6141 3070,7700" style="stroke:#000000;stroke-width:300; stroke-linejoin:miter; stroke-linecap:round;"></polyline><polyline points="6418,4582 7600,6141 6418,7700" style="stroke:#000000;stroke-width:300; stroke-linejoin:miter; stroke-linecap:round;"></polyline><polyline points="8976,3590 8976,9732" style="stroke:#000000;stroke-width:174; stroke-linejoin:miter; stroke-linecap:round;"></polyline><path style="stroke:#000000;stroke-width:174;stroke-linecap:butt;" d="M 4960,9685 A 4228 4228  0  0  1  8976  9732"></path></g></svg></div></a><a class="runestone-link" href="https://runestone.academy" title="Runestone Academy"><img class="logo" src="https://runestone.academy/runestone/static/images/RAIcon_cropped.png"></a><a class="mathjax-link" href="https://www.mathjax.org" title="MathJax"><img class="logo" src="https://www.mathjax.org/badge/badge-square-2.png"></a>
</div>
</body>
</html>
